# Assignment 1

Files and CSV

 The textfile, `assets/travel_plans.txt`, contains the summer travel plans for someone with some commentary. Find the total number of characters in the file and save to the variable `num`.

<!-- grade_id: cell-774c27f815c07a1e, points: 0, grade: false, solution: true -->
```python
num = None

# YOUR CODE HERE
import csv
file= open('assets/travel_plans.txt',"r")
num=0

for line in file:
    print(line)
    num=num+len(line)
    print(num)
raise NotImplementedError()
```

Output:
```
This summer I will be travelling.

34
I will go to...

50
Italy: Rome

62
Greece: Athens

77
England: London, Manchester

105
France: Paris, Nice, Lyon

131
Spain: Madrid, Barcelona, Granada

165
Austria: Vienna

181
I will probably not even want to come back!

225
However, I wonder how I will get by with all the different languages.

295
I only know English!

316
NotImplementedError                       Traceback (most recent call last)
Cell In[40], line 12
     10     num=num+len(line)
     11     print(num)
---> 12 raise NotImplementedError()

NotImplementedError:
```

Sometimes, as in the next cell, we will show you one or more test cases, in the form of assert statements. Make sure your code pass those tests before you click the "Submit Assignment" button.

When you click the "Submit Assignment" button, there may be additional, hidden tests that are run. If you fail one of those tests, you will get feedback and can resubmit, as many times as you want, until you get them all right.

But it takes a little while for the auto grader to run. You will find it is faster, less frustrating, and more educational if you create additional tests of your own, either with assert statements or just print statements. That will give you a faster feedback-and-fix cycle; and formulating those tests will help you really understand what's going on. 

<!-- grade_id: cell-05ae8996efdc7e7b, points: 1, grade: true, solution: false -->
```python
assert num == 316, "num is not assigned the correct value"

```

We have provided a file called `assets/emotion_words.txt` that contains lines of words that describe emotions. Find the total number of words in the file and assign this value to the variable `num_words`.

<!-- grade_id: cell-0e1c86fffdf89e01, points: 0, grade: false, solution: true -->
```python
num_words = None

# YOUR CODE HERE

with open("assets/emotion_words.txt","r") as file:
    lines=file.readlines()
    num_words=0
    for line in lines:
        words=line.split()
        num_words=num_words+len(words)
    
print(num_words)
    
    
raise NotImplementedError()
```

Output:
```
48
NotImplementedError                       Traceback (most recent call last)
Cell In[9], line 15
     10         num_words=num_words+len(words)
     12 print(num_words)
---> 15 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-7897d83bc4e61c17, points: 1, grade: true, solution: false -->
```python
assert num_words == 48, "num_words is not assigned the correct value"
```

Assign to the variable `num_lines` the number of lines in the file `assets/school_prompt.txt`.

**Note: different file than the previous question!**

<!-- grade_id: cell-75b61aedfeb412a9, points: 0, grade: false, solution: true -->
```python
num_lines = None

# YOUR CODE HERE
with open("assets/school_prompt.txt",'r') as file:
    lines=file.readlines()
    print(len(lines))
    num_lines=len(lines)
raise NotImplementedError()
```

Output:
```
10
NotImplementedError                       Traceback (most recent call last)
Cell In[15], line 8
      6     print(len(lines))
      7     num_lines=len(lines)
----> 8 raise NotImplementedError()

NotImplementedError:
```

```python
num_lines
```

<!-- grade_id: cell-9122382cacdf4dda, points: 1, grade: true, solution: false -->
```python
assert num_lines == 10, "num_lines is not assigned the correct value"
```

Assign the first 30 characters of `assets/school_prompt.txt` as a string to the variable `beginning_chars`.

<!-- grade_id: cell-4836ee01c51ef411, points: 0, grade: false, solution: true -->
```python
beginning_chars = None

# YOUR CODE HERE
with open("assets/school_prompt.txt","r") as file:
    beginning_chars=" "
    characters=file.read()
    beginning_chars=characters[0:30]
    print(len(beginning_chars))
    print(beginning_chars)
raise NotImplementedError()
```

Output:
```
30
Writing essays for school can 
NotImplementedError                       Traceback (most recent call last)
Cell In[29], line 10
      8     print(len(beginning_chars))
      9     print(beginning_chars)
---> 10 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-dae2a8b130911b15, points: 1, grade: true, solution: false -->
```python
assert type(beginning_chars) == type(""), "beginning_chars is not a string"
assert len(beginning_chars) == 30, "beginning_chars is not assigned the correct length"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[21], line 1
----> 1 assert type(beginning_chars) == type(""), "beginning_chars is not a string"
      2 assert len(beginning_chars) == 30, "beginning_chars is not assigned the correct length"

AssertionError: beginning_chars is not a string
```

**Challenge:** Using the file `assets/school_prompt.txt`, assign the third word of every line to a list called `three`.

<!-- grade_id: cell-3e6fa54e4fe79a87, points: 0, grade: false, solution: true -->
```python
three = None

# YOUR CODE HERE
with open("assets/school_prompt.txt","r") as file:
    x=0
    three=[]
    lines=file.readlines()
    for line in lines:
        line_words=line.split()
        three.append(line_words[2])
        print(line_words)
        print(len(line_words))
print(three)    
print(len(three))    
    
raise NotImplementedError()
```

Output:
```
['Writing', 'essays', 'for', 'school', 'can', 'be', 'difficult', 'but']
8
['many', 'students', 'find', 'that', 'by', 'researching', 'their', 'topic', 'that', 'they']
10
['have', 'more', 'to', 'say', 'and', 'are', 'better', 'informed.', 'Here', 'are', 'the', 'university']
12
['we', 'require', 'many', 'undergraduate', 'students', 'to', 'take', 'a', 'first', 'year', 'writing', 'requirement']
12
['so', 'that', 'they', 'can']
4
['have', 'a', 'solid', 'foundation', 'for', 'their', 'writing', 'skills.', 'This', 'comes']
10
['in', 'handy', 'for', 'many', 'students.']
5
['Different', 'schools', 'have', 'different', 'requirements,', 'but', 'everyone', 'uses']
8
['writing', 'at', 'some', 'point', 'in', 'their', 'academic', 'career,', 'be', 'it', 'essays,', 'research', 'papers,']
13
['technical', 'write', 'ups,', 'or', 'scripts.']
5
['for', 'find', 'to', 'many', 'they', 'solid', 'for', 'have', 'some', 'ups,']
10
NotImplementedError                       Traceback (most recent call last)
Cell In[26
```

<!-- grade_id: cell-d7416a181aa2a896, points: 1, grade: true, solution: false -->
```python
assert type(three) == type([]), "three is not a list"
assert len(three) == 10, "three is not assigned the correct length"
```

**Challenge:** Create a list called `emotions` that contains the first word of every line in `assets/emotion_words.txt`.

<!-- grade_id: cell-cc5eb2cda7bc9765, points: 0, grade: false, solution: true -->
```python
emotions = None

# YOUR CODE HERE
with open("assets/emotion_words.txt","r") as file:
    emotions=[]
    emotion_line_list=file.readlines()
    for line in emotion_line_list:
        print(line)
        split_word=line.split()
        
        emotions.append(split_word[0])
        print(emotions)
        
raise NotImplementedError()

```

Output:
```
Sad upset blue down melancholy somber bitter troubled

['Sad']
Angry mad enraged irate irritable wrathful outraged infuriated

['Sad', 'Angry']
Happy cheerful content elated joyous delighted lively glad

['Sad', 'Angry', 'Happy']
Confused disoriented puzzled perplexed dazed befuddled

['Sad', 'Angry', 'Happy', 'Confused']
Excited eager thrilled delighted

['Sad', 'Angry', 'Happy', 'Confused', 'Excited']
Scared afraid fearful panicked terrified petrified startled

['Sad', 'Angry', 'Happy', 'Confused', 'Excited', 'Scared']
Nervous anxious jittery jumpy tense uneasy apprehensive

['Sad', 'Angry', 'Happy', 'Confused', 'Excited', 'Scared', 'Nervous']
NotImplementedError                       Traceback (most recent call last)
Cell In[16], line 14
     11         emotions.append(split_word[0])
     12         print(emotions)
---> 14 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-cfad5425c19f0642, points: 1, grade: true, solution: false -->
```python
assert type(emotions) == type([]), "emotions is not a list"
assert len(emotions) == 7, "emotions is not assigned the correct length"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[28], line 1
----> 1 assert type(emotions) == type([]), "emotions is not a list"
      2 assert len(emotions) == 7, "emotions is not assigned the correct length"

AssertionError: emotions is not a list
```

Assign the first 33 characters from the textfile, `assets/travel_plans.txt` to the variable `first_chars`.

<!-- grade_id: cell-fdaec3f02923071d, points: 0, grade: false, solution: true -->
```python
first_chars = None

# YOUR CODE HERE
with open("assets/travel_plans.txt","r") as file:
    
    words=file.read()
    first_chars=words[0:33]
    print(type(first_chars))
raise NotImplementedError()
```

Output:
```
<class 'str'>
NotImplementedError                       Traceback (most recent call last)
Cell In[20], line 9
      7     first_chars=words[0:33]
      8     print(type(first_chars))
----> 9 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-29e19e35ea6223bb, points: 1, grade: true, solution: false -->
```python
assert type(first_chars) == type(""), "first_chars is not a string"
assert len(first_chars) == 33, "first_chars is not assigned the correct length"
```

**Challenge:** Using the file `assets/school_prompt.txt`, if the character ‘p’ is in a word, then add the word to a list called `p_words`.

<!-- grade_id: cell-20166062c59b6813, points: 0, grade: false, solution: true -->
```python
p_words = None

# YOUR CODE HERE
with open("assets/school_prompt.txt","r") as file:
    p_words=[]
    words=file.read()
    split_words=words.split()
    print(split_words)
    for word in split_words:
        if "p" in word:
            p_words.append(word)
            print(p_words)
raise NotImplementedError()
```

Output:
```
['Writing', 'essays', 'for', 'school', 'can', 'be', 'difficult', 'but', 'many', 'students', 'find', 'that', 'by', 'researching', 'their', 'topic', 'that', 'they', 'have', 'more', 'to', 'say', 'and', 'are', 'better', 'informed.', 'Here', 'are', 'the', 'university', 'we', 'require', 'many', 'undergraduate', 'students', 'to', 'take', 'a', 'first', 'year', 'writing', 'requirement', 'so', 'that', 'they', 'can', 'have', 'a', 'solid', 'foundation', 'for', 'their', 'writing', 'skills.', 'This', 'comes', 'in', 'handy', 'for', 'many', 'students.', 'Different', 'schools', 'have', 'different', 'requirements,', 'but', 'everyone', 'uses', 'writing', 'at', 'some', 'point', 'in', 'their', 'academic', 'career,', 'be', 'it', 'essays,', 'research', 'papers,', 'technical', 'write', 'ups,', 'or', 'scripts.']
['topic']
['topic', 'point']
['topic', 'point', 'papers,']
['topic', 'point', 'papers,', 'ups,']
['topic', 'point', 'papers,', 'ups,', 'scripts.']
NotImplementedError                       Traceback (m
```

<!-- grade_id: cell-fcfd04c36cdaf0a5, points: 1, grade: true, solution: false -->
```python
assert type(p_words) == type([]), "p_words is not a list"
assert len(p_words) == 5, "p_words is not assigned the correct length"
```